```python
# Initialize Otter
import otter
grader = otter.Notebook("mauna_loa.ipynb")
```

## Homework 8: Question 2
# 52 pts, Due Tuesday March 8 at 11:59pm 

```python
import pandas as pd
import matplotlib.pyplot as plt
import matplotlib
import numpy as np
%matplotlib inline
```

## Point Breakdown

| **Question #** | **Points**  |
| -------------- | :---------: |
| Question 0     | 4           |
| Question 1     | 35          |
| Question 2     | 43          |

## Question 1: Making a gh-pages site (8 pts)
Instructions to help you through this question can be found on [the class web site](https://bds311.github.io/gh_pages).

Please post below a URL to your gh-pages repository (i.e. `https://github.com/github-username/github-username.github.io`), created with gh-pages according to the assignment directions. **Reminder** - the repository should have:
1. An `index.md` file that is the homepage for your website
2. A second `experience.md` in your repository
3. A link in `index.md` that navigates website visitors to the page outlined by your second markdown file.





_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Question 2a (6 pts):
Read `mauna_loa.csv` into Pandas as a dataframe and make a scatter plot of C02 over time using the `C02` and `decimal_date` columns.

<!-- BEGIN QUESTION -->



```python
...
```

<!-- END QUESTION -->

### Question 2b (12 pts):
In `regression_fxns.py` file, complete the following three provided function definitions:

```
def make_standard_units(input_array):
    Converts input_array to standard_units, where data has mean 0 and standard deviation of 1
        INPUT: data array
        OUTPUT: array in standard units
    
def calc_corrcoef_from_standardized_input(array1,array2):
    Calculates Pearson correlation coefficient from two arrays in standard units
    INPUT: array1, array2: In standard units
    OUTPUT: Pearson correlation coefficient

def get_regression_parameters(array1, array2):
    Calculates regression parameters from two input arrays
    INPUT: array1, array2: two data arrays
    OUTPUT: regression_array, length 2: regression_array[0] is slope and regression_array[1] is intercept'''
```
The import regression_fxns below and use these methods to plot a predicted linear regression for the data above. Extend your fit line to 2100.

Remember that the slope of the regression line is `correlation_coefficient * (sd(y)/sd(x))`.
The intercept can be calculated from the equation of a line using the mean of the data and the slopes `meany = slope * meanx + intercept`

**NOTE**: If you import your `regression_fxns.py` file and then make changes to the file, <u><em>you must restart your kernel (and re-import the file) for the changes to take effect</em></u>

<!-- BEGIN QUESTION -->



```python
# Import regression functions
...

# User the functions in regression_fxns.py to graph a line of best fit on the data
# Remember - line of best fit should be extended to 2100
...
```

<!-- END QUESTION -->

### Question 2c (6 pts):
Based on calculated best-fit line, what would you predict the CO2 level to be in ppm in 2050, in 2100?

**Note**: Autograder values used here are not the correct answers, only close to the correct answers. These autograder values should be used as an indicator that you're on the right track, not a 100% guarantee that you've gotten the question right.

```python
co2_2050 = ...
print(co2_2050)
co2_2100= ...
print(co2_2100)
```

```python
grader.check("q2c")
```

### Question 2d:
If you were to fit only based on last 10 years of recorded data, would predicted CO2 levels be higher or lower? Why?

_Type your answer here, replacing this text._

<!-- END QUESTION -->



---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

```python

```

# Special Submission Instructions for this assignment:
### Rather than uploading .zip file to gradescope, upload your completed `mauna_loa.ipynb` and `regression_fxns.py` files in gradescope folder.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export()
```

Tests for q2c (points: 6)
```
>>> assert np.isclose(co2_2050, 452, rtol = 10)
```
```
>>> assert np.isclose(co2_2100, 531, rtol = 10)
```
